# BT02: Pandas

Họ tên: ...

MSSV: ...

---
## Cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT02`
        - File `BT02-Pandas.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác cấu trúc này.</font>

---
## Mô tả dữ liệu
Dữ liệu được dùng trong bài tập này là file "mvt.csv" đính kèm - đây là dữ liệu về tội phạm ăn cắp xe ở thành phố Chicago, bang Illinois, Hoa Kỳ ("mvt" là viết tắt của "motor vehicle theft"). Dữ liệu này được lấy từ khóa học "The Analytics Edge" ở trang edx.org, và gốc của dữ liệu là [ở đây](https://data.cityofchicago.org/Public-Safety/Crimes-2001-to-present/ijzp-q8t2).

Dưới đây là phần mô tả của khóa học "The Analytics Edge" về các cột trong file "mvt.csv":
- **ID**: a unique identifier for each observation.
- **Date**: the date the crime occurred.
- **LocationDescription**: the location where the crime occurred.
- **Arrest**: whether or not an arrest was made for the crime (TRUE if an arrest was made, and FALSE if an arrest was not made).
- **Domestic**: whether or not the crime was a domestic crime, meaning that it was committed against a family member (TRUE if it was domestic, and FALSE if it was not domestic).
- **Beat**: the area, or "beat" in which the crime occurred. This is the smallest regional division defined by the Chicago police department.
- **District**: the police district in which the crime occured. Each district is composed of many beats, and are defined by the Chicago Police Department.
- **CommunityArea**: the community area in which the crime occurred. Since the 1920s, Chicago has been divided into what are called "community areas", of which there are now 77. The community areas were devised in an attempt to create socially homogeneous regions.
- **Year**: the year in which the crime occurred.
- **Latitude**: the latitude of the location at which the crime occurred.
- **Longitude**: the longitude of the location at which the crime occurred.

---
## Import

```python
%matplotlib inline
import pandas as pd
# YOUR CODE HERE (OPTION) 
# Nếu cần các thư viện khác thì bạn có thể import ở đây
```

---
## Đọc dữ liệu từ file csv vào dataframe

```python
mvt_df = pd.read_csv("mvt.csv")
mvt_df.head()
```

---
## 1. Xem số dòng, số cột, và các giá trị thiếu (1đ)
Nhìn vào kết quả của câu lệnh `mvt_df.info()` và cho biết:
- `mvt_df` có bao nhiêu dòng và bao nhiêu cột? Bạn chỉ cần điền cứng kết quả vào hai biến `num_rows` và `num_cols` ở dưới là được.
- `mvt_df` có những cột nào có giá trị thiếu (null/na) và có bao nhiêu giá trị thiếu? Bạn chỉ cần điền cứng kết quả vào dictionary `na_counts` ở dưới, trong đó: key là tên cột có giá trị thiếu, và value là số lượng giá trị thiếu tương ứng.

```python
mvt_df.info()
```

<!-- grade_id: cell-67ad53ac852a545e, points: 0, grade: false, solution: true -->
```python
num_rows = 0 # Số dòng
num_cols = 0 # Số cột
na_counts = {} # Dictionary chứa tên các cột có giá trị thiếu 
               # và số lượng giá trị thiếu tương ứng
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cb30f534afe14d28, points: 1, grade: true, solution: false -->
```python
# TEST
assert num_rows > 190000
assert num_cols > 10
assert len(na_counts) == 4
assert na_counts["District"] > 43000
assert na_counts["CommunityArea"] > 24000
assert na_counts["Latitude"] > 2000
assert na_counts["Longitude"] > 2000
```

---
## 2. Mở kiểu dữ liệu object (1đ)
Nhìn vào kết quả của câu lệnh `mvt_df.info()`, ta thấy có 4 cột có `dtype` là `object` (`dtype` ám chỉ kiểu dữ liệu của các phần tử trong cột, chứ không phải kiểu dữ liệu của cột, cột có kiểu dữ liệu là series): `Date`, `LocationDescription`, `District`, `CommunityArea`. Trong Pandas, kiểu dữ liệu `object` thường ám chỉ chuỗi, nhưng thật ra kiểu dữ liệu `object` có thể chứa một đối tượng bất kỳ trong Python (vì thật ra ở bên dưới kiểu dữ liệu `object` chứa địa chỉ). Nếu một cột trong dataframe có `dtype` là `object` thì có thể các phần tử trong cột này sẽ có kiểu dữ liệu khác nhau; để biết được kiểu dữ liệu thật sự của các phần tử trong cột này thì ta phải truy xuất vào từng phần tử và dùng câu lệnh `type` để xem kiểu dữ liệu, ví dụ `type(s.iloc[i])` với `s` là một cột. Lúc đọc dữ liệu từ file "mvt.csv" lên, nếu bạn để ý thì sẽ thấy warning là: `Columns (6,7) have mixed types`, nghĩa là cột 6 (`District`) và 7 (`CommunityArea`) có `dtype` không đồng nhất (trong nội bộ mỗi cột). Ta muốn xem thử trong nội bộ mỗi cột này có các kiểu dữ liệu nào.

Công việc của bạn là viết hàm nhận đầu vào là một cột (series) có `dtype` là `object` và trả về tập các kiểu dữ liệu thật sự có trong cột này.

<!-- grade_id: cell-bff98f5183b2dd1c, points: 0, grade: false, solution: true -->
```python
def open_object_dtype(s):
    """
    Tìm tập các kiểu dữ liệu có trong series s (s có dtype là object).
    """
    dtypes = set()
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return dtypes
```

<!-- grade_id: cell-c74acde905565cc8, points: 1, grade: true, solution: false -->
```python
# TEST
assert open_object_dtype(mvt_df.Date) == {str}
assert open_object_dtype(mvt_df.LocationDescription) == {str}
assert open_object_dtype(mvt_df.District) == {str, float}
assert open_object_dtype(mvt_df.CommunityArea) == {str, float}
```

```python
# float ở 2 cột District và CommunityArea chính là nan (giá trị thiếu)
print(mvt_df.District[36]) # Một giá trị nan trong cột District
print(type(mvt_df.District[36]))
```

---
## 3. Xử lý kiểu dữ liệu thời gian (1đ)
Hiện giờ, cột `Date` có `dtype` là `object` (`str`). Công việc tiếp theo của bạn là:
1. Chuyển `dtype` của cột `Date` sang kiểu dữ liệu thời gian bằng câu lệnh `pd.to_datetime` (sau khi làm xong thì cột `Date` trong `mvt_df` phải thay đổi `dtype`). Khi dùng `pd.to_datetime`, bạn truyền tham số `format="%m/%d/%y %H:%M"` (nếu không thì sẽ chạy rất chậm vì Pandas phải đoán format).
2. Thêm cột `Month` - tháng trong cột `Date` - vào `mvt_df`.
3. Thêm cột `WeekDay` - ngày trong tuần trong cột `Date` - vào `mvt_df` (0 là thứ hai, 1 là thứ ba, ..., 6 là chủ nhật).
4. Dùng phương thức `min` và `max` của cột `Date`, cho biết dữ liệu là từ thời điểm nào đến thời điểm nào? Bạn lưu kết quả vào hai biến `min_date` và `max_date`.

<!-- grade_id: cell-4471ff481055fed4, points: 0, grade: false, solution: true -->
```python
min_date = pd.Timestamp("")
max_date = pd.Timestamp("")
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43bfb9a4d84b1507, points: 1, grade: true, solution: false -->
```python
# TEST
assert all(mvt_df.Month == mvt_df.Date.dt.month)
assert all(mvt_df.WeekDay == mvt_df.Date.dt.dayofweek)
assert min_date == pd.Timestamp("2001-01-01 00:01:00")
assert max_date == pd.Timestamp("2012-12-31 23:15:00")
```

---
## 4. Vẽ đồ thị 1 (1đ)
Để có cảm nhận sơ bộ về sự thay đổi số lượng các vụ phạm tội theo thời gian, bạn hãy vẽ histogram (dạng không chuẩn hóa) của cột `Date` bằng phương thức `hist` của cột `Date` với số lượng bin là 50. Bạn đặt tên trục hoành là `Date` và tên trục tung là `Num crimes`. 

<!-- grade_id: cell-33e173b1ceb218c6, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## 5. Gom nhóm 1 (1đ)
Cho biết (dùng `groupby`): 
- Năm nào có số lượng vụ phạm tội nhiều nhất? Bạn lưu kết quả vào biến `bad_year`.
- Tháng nào có số lượng vụ phạm tội nhiều nhất? Bạn lưu kết quả vào biến `bad_month`.
- Ngày nào trong tuần có số lượng vụ phạm tội nhiều nhất? Bạn lưu kết quả vào biến `bad_week_day`.

<!-- grade_id: cell-5978daf2f1019f91, points: 0, grade: false, solution: true -->
```python
bad_year = 0; bad_month = 0; bad_weed_day = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-246c3835cb636a60, points: 1, grade: true, solution: false -->
```python
# TEST
assert bad_year == 2001
assert bad_month == 10
assert bad_week_day == 4
```

---
## 6. Gom nhóm 2 (1đ)
Cho biết tỉ lệ các vụ phạm tội mà đã bắt được thủ phạm trong mỗi năm (tỉ lệ của mỗi năm được tính bằng cách lấy <font color=blue>số vụ phạm tội trong năm đó mà đã bắt được thủ phạm</font> chia cho <font color=blue>tổng số vụ phạm tội trong năm đó</font>)? Bạn lưu kết quả vào series `arrest_proportions` ở dưới (với index là năm). Gợi ý: bạn có thể gom nhóm bằng list các tên cột, sau đó có thể dùng phương thức `unstack` để tổ chức lại dữ liệu và dễ dàng tính toán.

<!-- grade_id: cell-c44b745e9b5ecde6, points: 0, grade: false, solution: true -->
```python
arrest_proportions = pd.Series()
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-faee661af8250853, points: 1, grade: true, solution: false -->
```python
# TEST
assert f"{arrest_proportions.loc[2001]:.3f}" == "0.104"
assert f"{arrest_proportions.loc[2006]:.3f}" == "0.081"
assert f"{arrest_proportions.loc[2012]:.3f}" == "0.039"
```

---
## 7. Vẽ đồ thị 2 (1đ)
Sau khi tính được series `arrest_proportions` ở trên, để có cảm nhận rõ hơn về sự thay đổi tỉ lệ các vụ phạm tội mà đã bắt được thủ phạm qua các năm, bạn dùng phương thức `plot` của series `arrest_proportions` để vẽ đồ thị dạng line.
- Trục hoành là năm. Bạn đặt tên là `Year`.
- Trục tung là tỉ lệ bị bắt. Bạn đặt tên là `Arrest proportion`. Bạn cũng cần phải chỉnh lại miền giá trị của trục tung là từ 0 đến 1.

<!-- grade_id: cell-84eac7aaa03790b3, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Bạn sẽ thấy tỉ lệ bị bắt nhìn chung là giảm trong các năm gần đây. Tại sao lại như vậy? Về cơ bản là ta không biết chắc được. Nhưng bạn thử suy nghĩ và đưa ra một giả thuyết để giải thích cho xu hướng này dựa vào thông tin sau: dữ liệu này được thu thập vào khoảng năm 2013-2014 (thời điểm ra đời của khóa học "The Analytics Edge").

<!-- grade_id: cell-b5fb9a5a4169d994, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

---
## 8. Đếm số lần xuất hiện của giá trị (1đ)
Cho biết 5 loại địa điểm (dữ liệu về loại địa điểm lưu ở cột `LocationDescription`) mà xảy ra nhiều vụ phạm tội nhất (không tính loại địa điểm `OTHER`)? Bạn lưu kết quả vào list `top5_locations` ở dưới.

<!-- grade_id: cell-2d23e94fd354ea7a, points: 0, grade: false, solution: true -->
```python
top5_locations = [] # List
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9dcad232b3c5a2f4, points: 1, grade: true, solution: false -->
```python
# TEST 
assert set(top5_locations) == set(['STREET', 'PARKING LOT/GARAGE(NON.RESID.)', 'ALLEY', 'GAS STATION', 'DRIVEWAY - RESIDENTIAL'])
```

---
## 9. Lọc các dòng theo điều kiện (1đ)
Tìm các dòng trong `mvt_df` có giá trị `LocationDescription` thuộc `top5_locations` ở câu trước. Bạn lưu kết quả vào dataframe `top5_locations_df`.

<!-- grade_id: cell-aa0201e7806cc65f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-729cd3a97ce3e347, points: 1, grade: true, solution: false -->
```python
# Test
assert top5_locations_df.shape == (177510, 13)
assert set(top5_locations_df.LocationDescription.unique()) == set(['STREET', 'PARKING LOT/GARAGE(NON.RESID.)', 'ALLEY', 'GAS STATION', 'DRIVEWAY - RESIDENTIAL'])
```

---
## 10. Vẽ đồ thị 3 (1đ)
Từ `top5_locations_df` ở trên, vẽ đồ thị dạng cột (bar chart), trong đó:  trục hoành là loại địa điểm, trục tung là số lượng vụ phạm tội. Như vậy, mỗi cột sẽ ứng với một loại địa điểm, và chiều cao của cột ứng với số lượng vụ phạm tội tại loại địa điểm đó. Tuy nhiên, trong mỗi cột phải chia làm 2 phần (dùng 2 màu khác nhau): một phần ứng với số lượng các vụ án mà đã bắt được thủ phạm, và một phần ứng với số lượng các vụ án mà chưa bắt được thủ phạm. Bạn xem thêm hình mẫu bên dưới để hiểu cần vẽ như thế nào (không cần chính xác 100% như hình mẫu; ví dụ, màu sắc có thể khác).

<img src="BT02-barchart.png"/>

<!-- grade_id: cell-b3ecdeb359aee8c3, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```